**Homework 2**

In this homework you will practice using Python for file text processing and data list transformations. Make sure to read the HW requirements at the end of the notebook first. The 5 problems build on each other, so its OK to use results from a previous problem in a later problem.

**Enter your name:**  Thomas Lu

**Enter your PID:**  thomaslu

I have neither given nor received unauthorized assistance on this assignment.

**Problem 1.** Using Python, read in the data in the given "counties.csv" file (download from the Canvas HW2 page). Represent the data table in the form of a python list of tuples, where each tuple represents a county and contains the 8 column values for that county. Use appropriate data types for each column.  Preserve the order of the counties and attributes from the file. Do not include the header line as a tuple in your answer. You are not allowed to use CSV parsing libraries, you must parse the file yourself using functions like .read(), .split(), .splitlines(), etc.

<!-- grade_id: Q1, points: 0, grade: false, solution: true -->
```python
# Problem 1:
Answer1 = []  # Do not edit this line.
file = open('counties.csv')  # Do not edit this line.

# YOUR CODE HERE
result = []
next(file)
for line in file:
    fields = line.split(',')
    fields[-1] = fields[-1].strip()
    result.append((str(fields[0]),str(fields[1]),int(fields[2]),float(fields[3]),int(fields[4]), int(fields[5]), int(fields[6]), int(fields[7])))
# raise NotImplementedError("Assign result to variable Answer 1")
# Assign your result to variable Answer1.
Answer1 = result
tuple(Answer1)

print(Answer1[:5], Answer1[-5:])  # Do not edit this line.
```

Output:
```
[('Aleutians East', 'AK', 2305, 12.9, 20114, 1410, 6693, 518), ('Aleutians West', 'AK', 5259, 14.8, 20298, 4321, 24233, 428), ('Anchorage', 'AK', 251335, 26.9, 26619, 93037, 2612640, 528), ('Bethel', 'AK', 15525, 13.1, 15493, 2269, 59687, 460), ('Bristol Bay', 'AK', 1023, 18.9, 29699, 392, 9653, 464)] [('Sweetwater', 'WY', 40976, 13.3, 19767, 14141, 343920, 279), ('Teton', 'WY', 13444, 30.0, 35983, 10359, 251764, 392), ('Uinta', 'WY', 20087, 14.3, 16105, 4928, 135950, 258), ('Washakie', 'WY', 8657, 18.4, 18342, 2598, 57057, 236), ('Weston', 'WY', 6506, 12.7, 20620, 1734, 39019, 217)]
```

<!-- grade_id: A1, points: 20, grade: true, solution: false -->
```python
# Autograder 
import nose.tools
nose.tools.assert_equal(len(Answer1), 3140)
nose.tools.assert_equal(Answer1[0], ('Aleutians East', 'AK', 2305, 12.9, 20114, 1410, 6693, 518))
nose.tools.assert_equal(Answer1[-1], ('Weston', 'WY', 6506, 12.7, 20620, 1734, 39019, 217))
# Hidden tests here
```

**Problem 2.** Which counties are most/least educated?  Sort the list of tuples from problem 1 according to the "PercentCollegeGrad" attribute in descending order.

<!-- grade_id: Q2, points: 0, grade: false, solution: true -->
```python
# Problem 2
Answer2 = []  # Do not edit this line.

# YOUR CODE HERE

result2 = sorted(Answer1, key=lambda x: x[3], reverse=True)
#raise NotImplementedError()
# Assign your result to variable Answer2.
Answer2 = result2
tuple(Answer2)

print(Answer2[:5], Answer2[-5:])  # Do not edit this line.
```

Output:
```
[('Los Alamos', 'NM', 18604, 53.4, 28753, 6369, 83922, 403), ('Falls Chruch', 'VA', 9617, 52.8, 0, 12052, 236246, 769), ('Arlington', 'VA', 172660, 52.3, 34216, 100896, 1588830, 678), ('Montgomery', 'MD', 809569, 49.9, 34299, 328052, 7396838, 698), ('Pitkin', 'CO', 13589, 49.8, 39481, 14749, 292946, 663)] [('Van Buren', 'TN', 5095, 4.1, 10718, 864, 6944, 132), ('Loving', 'TX', 140, 4.0, 29643, 0, 0, 175), ('Morgan', 'TN', 18280, 3.7, 11375, 1880, 33539, 165), ('Menominee', 'WI', 4601, 3.7, 0, 408, 4414, 103), ('Yellowstone National Park (Part)', 'MT', 54, 0.0, 0, 0, 0, 207)]
```

<!-- grade_id: A2, points: 20, grade: true, solution: false -->
```python
# Autograder 
nose.tools.assert_equal(len(Answer2), 3140)
nose.tools.assert_greater_equal(Answer2[0][3], Answer2[1][3])
# Hidden tests here
```

**Problem 3.** What is the mean and standard deviation of the education levels?  Compute a  pair (mu, signma) containing the mean (mu) and standard deviation (sigma) of the "PercentCollegeGrad" attribute over all counties. You are not allowed to use library functions for mean and standard deviation.  You must compute them yourself using the following definitions (N = number of tuples):
![formula](https://wikimedia.org/api/rest_v1/media/math/render/svg/32e3c0f27c2595926963cc5d8df113e6a12cf917)
https://wikimedia.org/api/rest_v1/media/math/render/svg/32e3c0f27c2595926963cc5d8df113e6a12cf917

<!-- grade_id: Q3, points: 0, grade: false, solution: true -->
```python
# Problem 3
from math import sqrt
Answer3 = (0.0,0.0)  # Do not edit this line.

# YOUR CODE HERE
result3 = []
percentGrad = []
sum3 = 0
file = open('counties.csv')
next(file)
for line in file:
    fields = line.split(',')
    fields[-1] = fields[-1].strip()
    percentGrad.append(float(fields[3]))
    sum3 = sum3 + float(fields[3])
mean3 = sum3/len(Answer1)
stddev = sqrt((sum((i - mean3) ** 2 for i in percentGrad))/len(Answer1))

result3.append(mean3)
result3.append(stddev)
#raise NotImplementedError()
# Assign your result to variable Answer3.
Answer3 = result3

print(Answer3)  # Do not edit this line.
```

Output:
```
[13.515859872611484, 6.5737072832103385]
```

<!-- grade_id: A3a, points: 10, grade: true, solution: false -->
```python
# Autograder 
nose.tools.assert_equal(len(Answer3), 2)
# Hidden tests here
```

<!-- grade_id: A3b, points: 10, grade: true, solution: false -->
```python
# Autograder 
# Hidden tests here
```

**Problem 4.** Which counties are outliers in education level?  Generate a list of counties that are outliers on the "PercentCollegeGrad" attribute, ordered according to the file, with tuples containing only the Name, State, PercentCollegeGrad attributes. Outliers are defined as values outside of:  mu +- 2*sigma .

<!-- grade_id: Q4, points: 0, grade: false, solution: true -->
```python
# Problem 4
Answer4 = []  # Do not edit this line.

# YOUR CODE HERE
result4 = []
outlier = []
file = open('counties.csv')
next(file)
for line in file:
    fields = line.split(',')
    fields[-1] = fields[-1].strip()
    if((float(fields[3]) > (mean3 + (2*stddev))) or (float(fields[3]) < (mean3 - (2*stddev)))):
        outlier.append((str(fields[0]), str(fields[1]), float(fields[3]))) 
#raise NotImplementedError()
# Assign your result to variable Answer4.
Answer4 = outlier
print(Answer4)  # Do not edit this line.
```

Output:
```
[('Anchorage', 'AK', 26.9), ('Juneau', 'AK', 30.7), ('Madison', 'AL', 30.1), ('Shelby', 'AL', 29.0), ('Alameda', 'CA', 28.8), ('Contra Costa', 'CA', 31.6), ('Marin', 'CA', 44.0), ('Orange', 'CA', 27.8), ('San Francisco', 'CA', 35.0), ('San Mateo', 'CA', 31.3), ('Santa Clara', 'CA', 32.6), ('Santa Cruz', 'CA', 29.7), ('Yolo', 'CA', 30.3), ('Arapahoe', 'CO', 35.2), ('Boulder', 'CO', 42.1), ('Clear Creek', 'CO', 31.2), ('Denver', 'CO', 29.0), ('Douglas', 'CO', 40.7), ('Eagle', 'CO', 33.0), ('Gilpin', 'CO', 29.5), ('Grand', 'CO', 30.2), ('Gunnison', 'CO', 36.9), ('Hinsdale', 'CO', 32.0), ('Jefferson', 'CO', 30.7), ('La Plata', 'CO', 28.1), ('Larimer', 'CO', 32.3), ('Ouray', 'CO', 27.9), ('Pitkin', 'CO', 49.8), ('Routt', 'CO', 34.7), ('San Miguel', 'CO', 40.3), ('Summit', 'CO', 39.7), ('Fairfield', 'CT', 34.2), ('Middlesex', 'CT', 28.2), ('Tolland', 'CT', 29.2), ('Washington', 'DC', 33.3), ('Alachua', 'FL', 34.6), ('Leon', 'FL', 37.1), ('Clarke', 'GA', 37.5), ('Cobb', 'GA', 33.0), ('De Kalb
```

<!-- grade_id: A4, points: 20, grade: true, solution: false -->
```python
# Autograder 
nose.tools.assert_equal(len(Answer4), 161)
nose.tools.assert_equal(Answer4[0], ('Anchorage', 'AK', 26.9))
# Hidden tests here
```

**Problem 5.** [Difficult] What are the total populations of each US state?  Generate a list of the total populations of each state. The list should contain tuples of the form (State, Population) and be sorted by increasing population. DC is considered a state. Hint: use Python's "set" data structure, and figure out how to filter out and sum each state's county populations. This can be solved in 1-3 lines of code.

<!-- grade_id: Q5, points: 0, grade: false, solution: true -->
```python
# Problem 5
Answer5 = []  # Do not edit this line.

# YOUR CODE HERE
states = []
pop = []
result5 = []
file = open('counties.csv')
next(file)
for line in file:
    fields = line.split(',')
    fields[-1] = fields[-1].strip()
    states.append(str(fields[1]))
    pop.append(int(fields[2]))
sum = 0
for i in range(len(states) + 1):
    if(i == 0):
        sum = sum + pop[0]
    elif(i == len(states)):
        result5.append((str(states[i-1]), sum))
    elif(states[i] == states[i-1]):
        sum = sum + pop[i]
    else:
        result5.append((str(states[i-1]), sum))
        sum = 0
        sum = sum + pop[i]
result5 = sorted(result5, key=lambda x: x[1], reverse=False)
#raise NotImplementedError()
# Assign your result to variable Answer5.
Answer5 = result5


print(Answer5)  # Do not edit this line.
```

Output:
```
[('WY', 480184), ('DC', 554256), ('VT', 584771), ('AK', 601739), ('ND', 641367), ('DE', 717197), ('SD', 729034), ('MT', 870281), ('RI', 989794), ('NH', 1148253), ('ID', 1163261), ('HI', 1186724), ('ME', 1241382), ('NV', 1530108), ('NE', 1637112), ('NM', 1685401), ('WV', 1828140), ('UT', 1951408), ('AR', 2483769), ('KS', 2565328), ('MS', 2697243), ('IA', 2841764), ('OR', 3140585), ('CT', 3274662), ('OK', 3277687), ('SC', 3673287), ('CO', 3746585), ('KY', 3860219), ('AZ', 4217940), ('AL', 4252982), ('LA', 4342334), ('MN', 4609548), ('MD', 5042438), ('WI', 5122871), ('TN', 5256051), ('MO', 5323523), ('WA', 5430940), ('IN', 5803471), ('MA', 6073550), ('VA', 6618358), ('NC', 7195138), ('GA', 7200882), ('NJ', 7945298), ('MI', 9549353), ('OH', 11150506), ('IL', 11829940), ('PA', 12071842), ('FL', 14165570), ('NY', 18136081), ('TX', 18723991), ('CA', 31589153)]
```

<!-- grade_id: A5, points: 20, grade: true, solution: false -->
```python
# Autograder 
nose.tools.assert_equal(len(Answer5), 51)
nose.tools.assert_equal(Answer5[0], ('WY', 480184))
# Hidden tests here
```

**Programming Requirements for this Homework:**
All problems must be coded in Python, using lists and tuples. Libraries such as Numpy and Pandas are not allowed. 
To solve the problems, use map, reduce, filter, sort, slice transforms, as we discussed in class, such as python List Comprehensions.

**Submission:**
Complete the "YOUR CODE HERE" sections, and make sure all cells run correctly, including the Autograder cells.
Submit your completed Python Notebook file (MyPID_Homework2.ipynb) on Canvas.


**Honor Code:** (from the syllabus on Canvas in the Files tab)
The Undergraduate Honor Code pledge that each member of the university community agrees to abide by states:
“As a Hokie, I will conduct myself with honor and integrity at all times. I will not lie,
cheat, or steal, nor will I accept the actions of those who do.”
Students enrolled in this course are responsible for abiding by the Honor Code. A student who has doubts about how the Honor Code applies to any assignment is responsible for obtaining specific guidance from the course instructor before submitting the assignment for evaluation. Ignorance of the rules does not exclude any member of the University community from the requirements and expectations of the Honor Code. For additional information about the Honor Code, please visit: https://www.honorsystem.vt.edu/

The Virginia Tech Honor Code applies to this course and will be strictly enforced. Unless otherwise indicated by the professor, all assignments and exams must be **completed strictly on an individual basis**. It is acceptable to discuss assignments with classmates in only a general way, and to help each other learn the concepts and programming languages. However, the student’s submitted work should represent his/her own ideas. Students are not allowed to copy answers, including lines of code, from each other.

As mentioned above, program assignments are also considered individual work that allows you to demonstrate the skills you have learned. You may freely offer and receive verbal assistance on how to use the programming language, what library classes or methods do, what errors mean, or how to interpret assignment instructions with your classmates. However, **you may not share specific lines of code used to solve assignment problems**. You may not allow others to view your solution code on your screen or other medium, and you may not view other students’ solution code on their screens.

Submitting an assignment in this class assumes agreement to the Virginia Tech honor pledge for assignments as follows: “I have neither given nor received unauthorized assistance on this assignment.”

If you have questions or are unclear about what constitutes academic misconduct on an assignment, please speak with the instructor. The Honor Code is taken seriously in this course.


```python
# Ungraded scratch space.  Do not create any new cells above this point.

```