# Homework 9

**Enter your Name:**  Thomas Lu
    
** Enter your PID:**  thomaslu
    
I have neither given nor received unauthorized assistance on this assignment.

```python
# Run this cell first.  Do NOT edit this cell.
Answer1 = Answer2 = Answer3 = Answer4 = None
import pandas
import numpy
import requests
import bs4
import nose.tools
```

**Problem 1.**  Its march madness!  Everyone likes to argue about which conference is best.  To shed some light, scrape data (using python and the Requests and BeautifulSoup libraries) from this website:
http://www.ncaa.com/rankings/basketball-men/d1/ncaa-mens-basketball-rpi

In Answer1, return a DataFrame of the extracted School, Conference, and Rank for all 351 schools, sorted by rank from best to worst.


<!-- grade_id: Problem1, points: 0, grade: false, solution: true -->
```python
# Problem 1
# YOUR CODE HERE
#raise NotImplementedError()
url = 'http://www.ncaa.com/rankings/basketball-men/d1/ncaa-mens-basketball-rpi'
page = requests.get(url)
soup = bs4.BeautifulSoup(page.content,'html.parser')
table = soup.find('table',{'class':'sticky'})
rows = table.findAll('tr')
data = []
tablesch, tableconf, tablerank = ([] for i in range(3))
for tr in rows:
    data.append(tr.findAll('td'))
for i in range(1,352):
    temp = data[i][2].text
    temp = temp.strip(";")
    if(temp == 'A&M-Corpus; Christi'):
        temp = 'A&M-Corpus Christi'
    tablesch.append(temp)
    tableconf.append(data[i][3].text)
    tablerank.append(int(data[i][0].text))
Answer1 = pandas.DataFrame({
    "School": tablesch,
    "Conference": tableconf,
    "Rank": tablerank
})
Answer1
```

Output:
```
School    Conference  Rank
0             Villanova      Big East     1
1              Virginia           ACC     2
2                Kansas        Big 12     3
3                Xavier      Big East     4
4                  Duke           ACC     5
5        North Carolina           ACC     6
6            Cincinnati           AAC     7
7              Michigan       Big Ten     8
8                Purdue       Big Ten     9
9             Tennessee           SEC    10
10              Clemson           ACC    11
11             Kentucky           SEC    12
12               Nevada           MWC    13
13       Loyola Chicago           MVC    14
14         Michigan St.       Big Ten    15
15               Auburn           SEC    16
16              Arizona        Pac-12    17
17           Texas Tech        Big 12    18
18             Ohio St.       Big Ten    19
19              Houston           AAC    20
20          Wichita St.           AAC    21
21              Gonzaga           WCC    22
22   
```

<!-- grade_id: Grade1, points: 25, grade: true, solution: false -->
```python
# Problem 1 autograder -- run me
nose.tools.assert_is_instance(Answer1, pandas.DataFrame)
nose.tools.assert_tuple_equal(Answer1.shape, (351, 3))
nose.tools.assert_list_equal(list(Answer1.columns), ['School','Conference','Rank'])
nose.tools.assert_true(Answer1.Rank.is_monotonic_increasing)
# autograder hidden tests
'passed'
```

Output:
```
'passed'
```

**Problem 2.**  Which conference is best?  To settle the issue, use the data from Problem 1 to compute the mean Rank for each Conference.

In Answer2, return a Series of the mean Ranks, indexed by Conference, sorted best to worst.


<!-- grade_id: Problem2, points: 0, grade: false, solution: true -->
```python
# Problem 2
# YOUR CODE HERE
#raise NotImplementedError()
conferences = []
for i in Answer1['Conference']:
    if i not in conferences:
        conferences.append(i)
testvar2 = Answer1.groupby('Conference').mean()
testvar3 = pandas.Series(testvar2['Rank'],index=conferences)
Answer2 = testvar3.sort_values(ascending=True)
Answer2
```

Output:
```
Big 12            51.300000
SEC               58.500000
Big East          66.200000
ACC               68.000000
Pac-12            87.583333
Big Ten           94.214286
AAC              113.250000
MVC              124.300000
MWC              140.545455
MAC              153.500000
Atlantic 10      161.785714
CAA              168.800000
WCC              171.500000
C-USA            187.357143
WAC              195.125000
SoCon            197.300000
Big Sky          198.833333
Big West         204.111111
OVC              205.416667
Sun Belt         206.583333
MAAC             206.727273
Summit League    207.750000
America East     221.333333
Big South        230.000000
Ivy League       231.125000
Horizon          234.300000
Patriot          242.200000
NEC              243.200000
Southland        246.230769
ASUN             253.250000
MEAC             299.000000
SWAC             302.200000
Name: Rank, dtype: float64
```

<!-- grade_id: Grade2, points: 25, grade: true, solution: false -->
```python
# Problem 2 autograder -- run me
nose.tools.assert_is_instance(Answer2, pandas.Series)
nose.tools.assert_tuple_equal(Answer2.shape, (32, ))
nose.tools.assert_true(Answer2.is_monotonic_increasing)
nose.tools.assert_almost_equal(Answer2.iloc[0], 51.300000, delta=0.0001)
nose.tools.assert_equal(Answer2.index[0], 'Big 12')
# autograder hidden tests
'passed'
```

Output:
```
'passed'
```

**Problem 3.** 
Web-scraping is especially useful for integrating data together from many web pages.  This page has links to each team's stats in 2018:
http://www.sports-reference.com/cbb/seasons/2018-school-stats.html

Scrape the team names and links.

In Answer3, return a list of tuples containing each School name and its URL, in the same list order as on the web page.  Use the relative URLs as given in the page.

<!-- grade_id: Problem3, points: 0, grade: false, solution: true -->
```python
# Problem 3

# YOUR CODE HERE
#raise NotImplementedError()
url = 'https://www.sports-reference.com/cbb/seasons/2018-school-stats.html'
page = requests.get(url)
soup = bs4.BeautifulSoup(page.content,'html.parser')
table = soup.find('table',{'class':'sortable'})
data = table.findAll('a')
links = []
schools3 = []
tempanswer = []
for i in data:
    if i.has_attr('href'):
        links.append(i.attrs['href'])
        schools3.append(i.text)
for j in range(len(schools3)):
    tempanswer.append((schools3[j],links[j]))
Answer3 = tempanswer
tuple(Answer3)
```

Output:
```
(('Abilene Christian', '/cbb/schools/abilene-christian/2018.html'),
 ('Air Force', '/cbb/schools/air-force/2018.html'),
 ('Akron', '/cbb/schools/akron/2018.html'),
 ('Alabama A&M', '/cbb/schools/alabama-am/2018.html'),
 ('Alabama-Birmingham', '/cbb/schools/alabama-birmingham/2018.html'),
 ('Alabama State', '/cbb/schools/alabama-state/2018.html'),
 ('Alabama', '/cbb/schools/alabama/2018.html'),
 ('Albany (NY)', '/cbb/schools/albany-ny/2018.html'),
 ('Alcorn State', '/cbb/schools/alcorn-state/2018.html'),
 ('American', '/cbb/schools/american/2018.html'),
 ('Appalachian State', '/cbb/schools/appalachian-state/2018.html'),
 ('Arizona State', '/cbb/schools/arizona-state/2018.html'),
 ('Arizona', '/cbb/schools/arizona/2018.html'),
 ('Little Rock', '/cbb/schools/arkansas-little-rock/2018.html'),
 ('Arkansas-Pine Bluff', '/cbb/schools/arkansas-pine-bluff/2018.html'),
 ('Arkansas State', '/cbb/schools/arkansas-state/2018.html'),
 ('Arkansas', '/cbb/schools/arkansas/2018.html'),
 ('Army', '/cbb/
```

<!-- grade_id: Grade3, points: 25, grade: true, solution: false -->
```python
# Problem 3 autograder -- run me
nose.tools.assert_is_instance(Answer3, list)
nose.tools.assert_tuple_equal(Answer3[0], ('Abilene Christian', '/cbb/schools/abilene-christian/2018.html'))
# autograder hidden tests
'passed'
```

Output:
```
'passed'
```

**Problem 4.** What is the median Assists Per Game (AST column in 'Per Game' table) among players for each of the first 30 teams?

    
Use the links from Problem 3 to scrape the AST for each team's players from the 'Per Game' table. 
Next, for each team, compute the median AST among all players.  

In Answer4, return a DataFrame containing each School name and its median AST, sorted from best to worst.

<!-- grade_id: Problem4, points: 0, grade: false, solution: true -->
```python
# Problem 4
# YOUR CODE HERE
#raise NotImplementedError()
url = 'https://www.sports-reference.com'
pages = [bs4.BeautifulSoup(requests.get(url+Answer3[i][1]).content, 'html.parser') \
        for i in range(30)]
tables = [pages[i].find('table',{'class':'sortable', 'id':'per_game'}) \
        for i in range(len(pages))]
datas = [tables[i].findAll('tr') \
        for i in range(len(tables))]
tempans = []
#each team will have a list, then will be appended to the big list
teamlist = []
for i in range(len(datas)):
    for j in range(1,len(datas[i])):
        tempvar = datas[i][j].find('td',{'data-stat':'ast_per_g'})
        tempvar2 = tempvar.text
        teamlist.append(float(tempvar2))
        #teamlist.append(datas[i][j].find('td',{'data-stat':'ast_per_g'}))
        if(j == (len(datas[i]) - 1)):
            teammedian = numpy.median(teamlist)
            tempans.append(teammedian)
            teamlist = [] 
schools4 = []  
torg = pandas.DataFrame({
    "School": schools3[:30],
    "AST": tempans
})
Answer4 = torg.sort_values("AST",ascending=False)
Answer4
```

Output:
```
School   AST
20           Ball State  1.20
4    Alabama-Birmingham  1.20
23      Bethune-Cookman  1.15
27    Boston University  1.10
28  Bowling Green State  1.00
9              American  0.90
21               Baylor  0.90
8          Alcorn State  0.90
22              Belmont  0.80
12              Arizona  0.80
3           Alabama A&M  0.80
6               Alabama  0.75
13          Little Rock  0.75
19          Austin Peay  0.70
7           Albany (NY)  0.70
1             Air Force  0.65
15       Arkansas State  0.65
0     Abilene Christian  0.65
24           Binghamton  0.60
25          Boise State  0.60
29              Bradley  0.60
10    Appalachian State  0.55
18               Auburn  0.50
16             Arkansas  0.50
14  Arkansas-Pine Bluff  0.50
11        Arizona State  0.45
26       Boston College  0.40
5         Alabama State  0.40
2                 Akron  0.40
17                 Army  0.30
```

<!-- grade_id: Grade4, points: 25, grade: true, solution: false -->
```python
# Problem 4 autograder -- run me
nose.tools.assert_is_instance(Answer4, pandas.DataFrame)
nose.tools.assert_tuple_equal(Answer4.shape, (30, 2))
nose.tools.assert_list_equal(list(Answer4.columns), ['School','AST'])
nose.tools.assert_true(Answer4.AST.is_monotonic_decreasing)
nose.tools.assert_equal(Answer4.School.iloc[14], 'Albany (NY)')
nose.tools.assert_almost_equal(Answer4.AST.iloc[14], 0.7, delta=0.0001)
# autograder hidden tests
'passed'
```

Output:
```
'passed'
```

**Problem 5.**  (just for fun)

Which school will win the NCAA tournament this year?  

*Enter answer here:*  Duke


```python
# scratch space

```